## Problem 1: Simple scatter plot of random points (3 points)

The aim of this task is to create a simple scatter plot using random input values and random colors!
Finally, you will save the plot as `.png` file that should look something like this: 
 
![Example figure](img/problem1_example.png)
 
Remember to also commit the image file into your GitHub repo!

**NOTE**: You colors may look different from the image above depending on the colormap you are using!

### Tips for completing this problem

- Use **exactly** the same variable names as in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.
- **Please do not**:

    - **Change the file names**. Do all of your editing in the provided `Exercise-7-problem-1.ipynb` file (this file).
    - **Copy/paste cells in this notebook**. We use an automated grading system that will fail if there are copies of code cells.
    - **Change the existing cell types**. You can add cells, but changing the cell types for existing cells (from code to markdown, for example) will also cause the automated grader to fail.

### Scores for this problem

**Your score on this problem will be based on following criteria:**

- Creating a DataFrame containing 1000 x, y, and color values
- Creating a scatter plot of these points as instructed
- Adding a title and axis labels to the plot
- Saving a copy of the plot as a PNG file and commiting it to your repo
- Including comments that explain what most lines in the code do

### Part 1 (0.5 points)

Create an new data frame called `data` and add 1000 random numbers (of type `float`) into a new column `x` and another 1000 random numbers into a new column `y`.

<!-- grade_id: problem_1_ag_p1_random_points, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_1_ag_p1_andom_points_test1, points: 0.5, grade: true, solution: false -->
```python
# Check your random values
print(data.head())

# Check that you have the correct number of rows
assert len(data) == 1000, "There should be 1000 rows of data."
```

### Part 2 (0.5 points)

Create a variable `colors` in which you should store 1000 random values for the colors so that we get a little sweeter looking plot.

The variable `colors` should be an array of random numbers, not a DataFrame column.

<!-- grade_id: problem_1_ag_p2_colors, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_1_ag_p2_colors_test1, points: 0.5, grade: true, solution: false -->
```python
# This test print should print out 10 first numbers in the variable colors
print(colors[0:10])

# Check that the length matches
assert len(colors) == 1000, "There should be 1000 random numbers for colors"
```

### Part 3 (2 points)

#### Part 3.1

Create a scatter plot of points with random colors.

Make the scatter plot using `DataFrame.plot()` with the following options:

- set the type of the plot using parameter `kind=scatter` (or use `DataFrame.plot.scatter()`)
- modify the size of a point using the parameter `s` (the example plot uses `s=50`)
- set the random colors for the points using the parameter `c`
- set the color scheme using the parameter `colormap`
    - The example image uses the `'rainbow'` colormap
    - you can also choose another color map from here: https://matplotlib.org/stable/tutorials/colors/colormaps.html
    - you can add a black outline around the points using the parameter `edgecolors`. The example uses `edgecolors='black'`.

#### Part 3.2

Add a title, x-label and y-label to your plot as shown in the example figure.

- create the variables `title`, `xlabel` and `ylabel` and store the title and axis label text in them
- use the pandas plotting parameters or `matplotlib.pyplot` methods to add the title and axis labels to the figure

**Feel free to change the colors and markers in the plot :)**

#### Part 3.3

Create a variable `outputfp` with the output filename `my_first_plot.png` and save your plot as a PNG file into the disk (note: you have to save the file in the same code cell where you do the plotting). 

<!-- grade_id: problem_1_mg_p3_scatter_plot, points: 1.5, grade: true, solution: true -->
```python
%matplotlib inline

# Plot a scatter plot
# YOUR CODE HERE
raise NotImplementedError()

# Add labels and title (if not done above when creating the plot)
# YOUR CODE HERE
raise NotImplementedError()

# Save the plot as a png file:
outputfp = None

# YOUR CODE HERE
raise NotImplementedError()

# This test print statement should print the output filename of your figure
print(f'Saved my first plot as {outputfp}')
```

<!-- grade_id: problem_1_ag_p3_outputfile_test, points: 0.5, grade: true, solution: false -->
```python
# Check that the file exists (also go and open the file to check that everything is ok!)
import os

assert os.path.exists(outputfp), "Can't find the output image."
```

Remember to commit your changes (including the image file) to your GitHub repo!

### Done!

Now you can move to [Problem 2](Exercise-7-problem-2.ipynb).